---
title       : Insert the chapter title here
description : Insert the chapter description here

--- type:NormalExercise lang:r xp:100 skills:1
## More movies

asdf

*** =instructions
- asdf

*** =pre_exercise_code
```{r}
library(sparklyr)
sc <- spark_connect(master = "local")
```

*** =sample_code
```{r}
# sc is available
```

*** =solution
```{r}
# sc is available

```

*** =sct
```{r}
success_msg("Good work!")
```
